# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
"""
 Find the index of the longest word in a list of words.

 parameters:
 -words (list): A list of words.

 Returns:
 - int: Index of the longest word in the list.
 """
def longest(words):
  #print(type(words))
  li=[len(i) for i in words]
  long_index=li.index(max(li))
  return (long_index)

### END SOLUTION
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
```

### 25.2 Farenheit to Celsius

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION
"""
   converts a temperature from Fahrenhit to celsius.

   Parameters:
   -temperature_in_Farenheit(float): Temperature value in Fahrenheit.

   Returns:
   -float: Temperature value converted to celsius.

   Formula:
   celsius = 5/9*(Fahrenheit - 32)
"""
def f_to_c(temperature_in_Farenheit):
  temperature_in_celsius = 5/9 *(temperature_in_Farenheit-32)
  return temperature_in_celsius

### END SOLUTION
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates. Your function should should take three parameters: a start date, an end date, and a diagnosis.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION
from datetime import date

def los(start_date, end_date, diagnosis):
    """
    Compute the length of stay in whole days based on admit date, discharge date, and diagnosis.

    :param start_date: Admission date (date object)
    :type start_date: date
    :param end_date: Discharge date (date object)
    :type end_date: date
    :param diagnosis: Primary diagnosis
    :type diagnosis: str
    :return: Length of stay in whole days (0 if diagnosis is "Observation")
    :rtype: int
    """
    if diagnosis == 'Observation':
        return 0  # Length of stay is 0 for Observation diagnosis

    # Calculate the length of stay in days
    length_of_stay = (end_date - start_date).days

    return length_of_stay


### END SOLUTION
```

```python
assert los(date(2019,1,3), date(2019,1,8), "COPD") == 5
assert los(date(2020,1,5), date(2020,1,8), "Observation") == 0
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
```

```python
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for encid, visit in encounters.items():
    print(los(visit['admit'],visit['discharge'],visit['diagnosis']))

# Your output should be
# 5
# 4
# 1
# 0
# 0
```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python
### BEGIN SOLUTION
"""
   calculate the average Length of Stay (LOS) for a set of medical encounters.
   Parameters:
   -encounters (dict): A dictionary where keys are encounter IDs and values are dictionaries
    containing details such as admit date, discharge date, and diagnosis.

   Returns:
   -float: The average Length of AStay in days, excluding encounters with a diagnosis of "observation"
    Returns 0 if there are no valid encounters.
"""
def calculate_los(admit_date, discharge_date):
    """
    Calculate the length of stay (LOS) in whole days.

    :param admit_date: Admission date (date object)
    :type admit_date: date
    :param discharge_date: Discharge date (date object)
    :type discharge_date: date
    :return: Length of stay in whole days
    :rtype: int
    """
    return (discharge_date - admit_date).days

def average_los(encounters):
    """
    Calculate the average length of stay (LOS) for encounters in the provided dictionary.

    :param encounters: Dictionary of encounters with admit, diagnosis, and discharge dates
    :type encounters: dict
    :return: Average length of stay in whole days (excluding encounters with LOS of 0)
    :rtype: float
    """
    total_los = 0
    count = 0

    for encounter_id, data in encounters.items():
        admit_date = data.get('admit')
        discharge_date = data.get('discharge')
        diagnosis = data.get('diagnosis')

        if admit_date and discharge_date and diagnosis != 'Observation':
            los = calculate_los(admit_date, discharge_date)

            if los > 0:
                total_los += los
                count += 1

    if count == 0:
        return 0  # To avoid division by zero

    average = total_los / count
    return average

```

```python
from datetime import date
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(round(average_los(encounters),2) == 3.33)
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION
"""
   converts a temperature from celsius to Fahrenheit.

   Parameters:
   -temperature_in_celsius (float): Temperature value in celsius.

   Returns:
   -float: Temperature value converted to Fahrenheit.

   Formula:
   Fahrenheit = (celsius * 9/5) + 32
"""
def c_to_f(temperature_in_celsius):
   temperature_in_Fahrenheit = (temperature_in_celsius * 9/5) + 32
   return temperature_in_Fahrenheit


### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -a -m "message"`)
  5. Push your changes (`git push`)
  
If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.