# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In this GitHub repository, there is a file called `allergies.json` that contains a list of patient allergies.  You will need to download this [file from here](https://raw.githubusercontent.com/paulboal/hds5210-2023/main/week06/allergies.json) and then upload it into Google Colab to run these examples. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE="allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python


def allergy_count(json_file):
    """
    This function counts the number of entries in a JSON file.

    parameters:
        json_file (str): The name of the JSON file you want to process.

    returns:
        int: The count of entries in the JSON file.

    exceptions:
        - FileNotFoundError: Raised if the specified file does not exist.
        - json.JSONDecodeError: Raised if there is a problem decoding the JSON data.
        - ValueError: Raised if the JSON data does not contain an 'entry' key or is not in dictionary format.

    example:
        To count the number of entries in a JSON file named 'allergies.json', you can use this function like this:
        >>> json_file = 'allergies.json'
        >>> entry_count = allergy_count(json_file)
        >>> print(f'Number of entries in {json_file}: {entry_count}')
    """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            if isinstance(data, dict) and 'entry' in data:
                return len(data['entry'])
            else:
                raise ValueError("The JSON data does not contain an 'entry' key or is not in dictionary format.")
    except FileNotFoundError:
        print(f"Error: The file '{json_file}' does not exist.")
    except json.JSONDecodeError as e:
        print(f"Error decoding JSON in '{json_file}': {e}")
    except Exception as e:
        print(f"An error occurred: {e}")
    return 0






```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python


def patient_count(json_file):
    """
    count the number of unique patients in a JSON structure containing allergy data.

    parameters:
        json_file (str): The name of the JSON file containing allergy data.

    returns:
        int: The count of unique patients found in the JSON structure.

    exceptions:
        - FileNotFoundError: Raised if the specified file does not exist.
        - json.JSONDecodeError: Raised if there is a problem decoding the JSON data.
        - ValueError: Raised if the JSON data does not contain an 'entry' key or is not in dictionary format.

    example:
        To count the number of unique patients in a JSON file named 'allergies.json', use this function as follows:
        >>> unique_patients_count = patient_count('allergies.json')
        >>> print(f'Number of unique patients: {unique_patients_count}')
    """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            if isinstance(data, dict) and 'entry' in data:
                unique_patients = set()
                for entry in data['entry']:
                    if 'resource' in entry and 'patient' in entry['resource']:
                        unique_patients.add(entry['resource']['patient']['display'])
                return len(unique_patients)
            else:
                raise ValueError("The JSON data does not contain an 'entry' key or is not in dictionary format.")
    except FileNotFoundError:
        print(f"Error: The file '{json_file}' does not exist.")
    except json.JSONDecodeError as e:
        print(f"Error decoding JSON in '{json_file}': {e}")
    except Exception as e:
        print(f"An error occurred: {e}")
    return 0

```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python


def allergy_per_patient(json_file):
    """
    This function helps us count the number of allergies for each patient from a JSON file containing allergy data.

    parameters:
        json_file (str): The filename of the JSON file with allergy data.

    returns:
        dict: A dictionary where the keys represent patient names, and the values are the counts of allergies for each patient.

    exceptions:
        - FileNotFoundError: If the specified file does not exist.
        - json.JSONDecodeError: If there is an issue decoding the JSON data.
        - ValueError: If the JSON data does not contain an 'entry' key or is not in dictionary format.

    example:
        To find out how many allergies each patient has in a JSON file named 'allergies.json', use this function like this:
        >>> allergy_counts = allergy_per_patient('allergies.json')
        >>> print(allergy_counts)
        {'Paul Boal': 1, 'Jason Argonaut': 3}
    """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            if isinstance(data, dict) and 'entry' in data:
                allergy_counts = {}
                for entry in data['entry']:
                    if 'resource' in entry and 'patient' in entry['resource']:
                        patient_name = entry['resource']['patient']['display']
                        allergy_counts[patient_name] = allergy_counts.get(patient_name, 0) + 1
                return allergy_counts
            else:
                raise ValueError("The JSON data does not contain an 'entry' key or is not in dictionary format.")
    except FileNotFoundError:
        print(f"Error: The file '{json_file}' does not exist.")
    except json.JSONDecodeError as e:
        print(f"Error decoding JSON in '{json_file}': {e}")
    except Exception as e:
        print(f"An error occurred: {e}")
    return {}

```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python


def allergy_list(json_file):
    """
    extracts and formats information about patient allergies and their reactions from a JSON structure containing allergy data.

    parameters:
        json_file (str): The filename of the JSON file containing allergy data.

    returns:
        list: A list of lists where each inner list contains patient name, allergy, and the first reaction.

    exceptions:
        - FileNotFoundError: Raised if the specified file does not exist.
        - json.JSONDecodeError: Raised if there is an issue decoding the JSON data.
        - ValueError: Raised if the JSON data does not contain an 'entry' key or is not in dictionary format.

    example:
        To extract and format patient names, allergies, and reactions from a JSON file named 'allergies.json', you can use this function like this:
        >>> allergies_info = allergy_list('allergies.json')
        >>> print(allergies_info)
        [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
         ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
         ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
         ['Paul Boal', 'PENICILLIN G', 'Bruising']]
    """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            if isinstance(data, dict) and 'entry' in data:
                allergy_info = []
                for entry in data['entry']:
                    if 'resource' in entry and 'patient' in entry['resource'] and 'substance' in entry['resource'] and 'reaction' in entry['resource']:
                        patient_name = entry['resource']['patient']['display']
                        allergy = entry['resource']['substance']['text']
                        if 'manifestation' in entry['resource']['reaction'][0]:
                            reaction = entry['resource']['reaction'][0]['manifestation'][0]['text']
                            allergy_info.append([patient_name, allergy, reaction])
                return allergy_info
            else:
                raise ValueError("The JSON data does not contain an 'entry' key or is not in dictionary format.")
    except FileNotFoundError:
        print(f"Error: The file '{json_file}' does not exist.")
    except json.JSONDecodeError as e:
        print(f"Error decoding JSON in '{json_file}': {e}")
    except Exception as e:
        print(f"An error occurred: {e}")
    return []

```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python


def allergy_reaction(json_file, patient, substance):
    """
    Find the potential reaction that a patient may experience when taking a specified substance based on allergy data from a JSON file.

    parameters:
        json_file (str): The name of the JSON file containing allergy data.
        patient (str): The patient's name.
        substance (str): The substance for which you want to find the reaction.

    returns:
        str or None: The expected reaction if the patient takes the specified substance, or None if the substance is not found in the allergy data.

    examples:
        To find the potential reaction for a specific patient and substance in a JSON file named 'allergies.json', use this function like this:
        >>> reaction = allergy_reaction('allergies.json', 'Jason Argonaut', 'PENICILLIN G')
        >>> print(reaction)
        'Hives'

    exceptions:
        - FileNotFoundError: If the specified JSON file does not exist.
        - json.JSONDecodeError: If there is an issue decoding the JSON data.
        - ValueError: If the JSON data format is incorrect.

    """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            if isinstance(data, dict) and 'entry' in data:
                for entry in data['entry']:
                    if 'resource' in entry and 'patient' in entry['resource'] and 'substance' in entry['resource'] and 'reaction' in entry['resource']:
                        if entry['resource']['patient']['display'] == patient and entry['resource']['substance']['text'] == substance:
                            reaction = entry['resource']['reaction'][0]['manifestation'][0]['text'] if 'manifestation' in entry['resource']['reaction'][0] else None
                            if reaction:
                                return reaction
                return None
            else:
                raise ValueError("The JSON data does not contain an 'entry' key or is not in dictionary format.")
    except FileNotFoundError:
        print(f"Error: The file '{json_file}' does not exist.")
    except json.JSONDecodeError as e:
        print(f"Error decoding JSON in '{json_file}': {e}")
    except Exception as e:
        print(f"An error occurred: {e}")
    return None

```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

